# Assignment - Python Programming

<b><div style="text-align: right">[POINTS: 40]</div></b>

Python is an easy to learn, powerful programming language with efficient high-level data structures and object-oriented programming. Python’s elegant syntax and dynamic typing, together with its interpreted nature, make it an ideal language for scripting and rapid application development in many areas on most platforms.

In this assignment, you will not do one, big thing, but you will write small snippets of code that solve little problems.

### Jupyter Notebook

What you are reading now is an example of a Jupyter Notebook. The basic concept is that of a "notebook" containing text and programming code. You can easily edit the notebook using your web browser, run the programs in the ipython server in the background and see the output of the programs within the notebook. This is a powerful paradigm that is well suited to machine learning research, particularly when collaborating with other people. For Reference:

- https://ipython.org/notebook.html
- https://jupyter.org/




```python
# This rectangular box is a "Cell" in Jupyter Notebook (click to edit)
# This line starting with # symbol is a comment
# Press run on the menu or CTRL + ENTER to run this cell
# or SHIFT + ENTER to run this cell and move to next cell
# to escape from the current cell press ESC and H to see shortcut keys available

var1 = 5
var2 = 6
print(var1 * var2)
```

Output:
```
30
```

If you need to, feel free to go through the videos and reading resources throughout this unit. [This introductory tutorial](https://docs.python.org/3/tutorial/introduction.html) at python.org put up by the Python Software Foundation gives a really good introduction to the language if you need a refresher. Also, some of the exercises in this assignment assume that you are familiar with built-in data types in Python, such as `dict`s and `list`. So, check out the [documentation for the standard types](https://docs.python.org/3/library/stdtypes.html) in Python if you feel like you need to.

Lets begin by checking the version of Python being used. The `sys` module contains values for variables that the Python interpreter uses or maintains, and `sys.version` contains the version of Python being used.

```python
import sys
print(sys.version)
```

Output:
```
3.6.8 |Anaconda, Inc.| (default, Dec 30 2018, 01:22:34) 
[GCC 7.3.0]
```



**Task:** Check if a Python shell is executing in 32bit or 64bit mode.

Hint: you can use `struct.calcsize(fmt)` to calculate size of an character and multiply it by 8 to find out whether the python shell is executing in 32bit or 64bit mode. We multiply by 8 because the size is returned in terms of bytes.


```python
import struct
print(struct.calcsize("P") * 8)
```

Output:
```
64
```

## Exercise 1: Variables

<b><div style="text-align: right">[POINTS: 1]</div></b>

In this exercise you will create variables and assign values to them. You will also print out their types using a function built into Python.

**Task:** Create two variables named  `num` & `word` and assign `10` and `'Ten'` to them respectively. Also print out their types using the built-in `type` function.

<!-- grade_id: cell-679e2a0aa6f59056, points: 0, grade: false, solution: true -->
```python
num = 10
word = 'Ten'

# YOUR CODE HERE
print(type(num))
print(type(word))

```

Output:
```
<class 'int'>
<class 'str'>
```

<!-- grade_id: cell-479098952234f278, points: 1, grade: true, solution: false -->
```python
assert num is not None
assert word is not None

```

## Exercise 2: Casting between types

<b><div style="text-align: right">[POINTS: 2]</div></b>

**Task:** Convert the string `'1234.56'` in the variable `number` into `float`and store it in a variable `float_number`.

<!-- grade_id: cell-527f2975881fc2be, points: 0, grade: false, solution: true -->
```python
number="1234.56"
float_number = None

# YOUR CODE HERE
float_number = float(number)

```

<!-- grade_id: cell-f73899a0d60ae0b1, points: 2, grade: true, solution: false -->
```python
assert float_number is not None

```

## Exercise 3: Range

<b><div style="text-align: right">[POINTS: 2]</div></b>

**Task:** Write a program to store even numbers between 0 to 100 in a variable named `even_numbers`.

Hint: Use `range()` method.

<!-- grade_id: cell-f8bec99f9a9d9f2b, points: 0, grade: false, solution: true -->
```python
even_numbers= list(range(0,100,2))
even_numbers
```

Output:
```
[0,
 2,
 4,
 6,
 8,
 10,
 12,
 14,
 16,
 18,
 20,
 22,
 24,
 26,
 28,
 30,
 32,
 34,
 36,
 38,
 40,
 42,
 44,
 46,
 48,
 50,
 52,
 54,
 56,
 58,
 60,
 62,
 64,
 66,
 68,
 70,
 72,
 74,
 76,
 78,
 80,
 82,
 84,
 86,
 88,
 90,
 92,
 94,
 96,
 98]
```

<!-- grade_id: cell-87a0409bb270fec2, points: 2, grade: true, solution: false -->
```python
assert even_numbers is not None
```

## Exercise 4: Lists

<b><div style="text-align: right">[POINTS: 2]</div></b>

**Task:** Find the sum of the 10 smallest numbers in the list `L` and store it in a variable named `sum_list`.

`L = [75, 28, 47, 51, 96, 76, 32, 57, 10, 10, 25, 62, 17,  9, 68, 65, -2, 54, 34, 41, 74, 83, 91, 59, -3, 88, 59, 91, 10, -9, 90, 46, 58, 25, 62, 32, 74,  1, 61, 43, 25, 62, -5, 49, 40, 44, 83, -5, 13, 35]`

Python has builtin ways and functions to index, sort and sum up the values in a sequence of values, so be sure to use them.

<!-- grade_id: cell-8e86d41b9021b492, points: 0, grade: false, solution: true -->
```python
L =[75, 28, 47, 51, 96, 76, 32, 57, 10, 10, 25, 62, 17,  9, 68, 65, -2,
       54, 34, 41, 74, 83, 91, 59, -3, 88, 59, 91, 10, -9, 90, 46, 58, 25,
       62, 32, 74,  1, 61, 43, 25, 62, -5, 49, 40, 44, 83, -5, 13, 35]

sum_list = None
# YOUR CODE HERE
sum_list = sum(sorted(L)[:10])
print(sum_list)
```

Output:
```
[-9, -5, -5, -3, -2, 1, 9, 10, 10, 10]
16
```

<!-- grade_id: cell-ef493df9ed09c4af, points: 2, grade: true, solution: false -->
```python
assert sum_list is not None
```

## Exercise 5: Fibonnaci Numbers

<b><div style="text-align: right">[POINTS: 3]</div></b>

**Task:** Write a function `fibbo` that takes in an integer `n` and returns a list of `n` Fibonnaci numbers for `n` greater than 2.

Hint: The first two Fibonnaci numbers are `[1,1]`.

<!-- grade_id: cell-47228f28e6074a86, points: 0, grade: false, solution: true -->
```python
def fibbo(n):
    # YOUR CODE HERE
    fibonacci_numbers = [1, 1]
    for i in range(1,n-1):
        fibonacci_numbers.append(fibonacci_numbers[-2] + fibonacci_numbers[-1])
    return(fibonacci_numbers)

    
```

<!-- grade_id: cell-2e89da0d14026ef0, points: 3, grade: true, solution: false -->
```python
assert(fibbo(5)==[1, 1, 2, 3, 5])
assert(fibbo(11)==[1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89])
```

## Exercise 6: Tuples

<b><div style="text-align: right">[POINTS: 2]</div></b>

**Task:** `T` is a nested tuple, use indexing to extract 'g' from `T`. Store the result in variable `thisisg`.

Use `T = ('a','b','c',('d','e'),('f',('f','g'),'k'),'p')`

<!-- grade_id: cell-061453beb6d902de, points: 0, grade: false, solution: true -->
```python
T = ('a','b','c',('d','e'),('f',('f','g'),'k'),'p')
thisisg = None
# YOUR CODE HERE
thisisg = T[4][1][1]

```

<!-- grade_id: cell-ff9b29c9050cb8f6, points: 2, grade: true, solution: false -->
```python
assert thisisg is not None
```

## Exercise 7: Dictionaries

<b><div style="text-align: right">[POINTS: 2]</div></b>

Mark went to the market to guy some vegetables. `price` is a dictionary that maps the price of a vegetable to its name.

**Task:** Write a function `totalcost` to find his total cost given his shopping list.

<!-- grade_id: cell-9dc208e27777652d, points: 0, grade: false, solution: true -->
```python
price = {"Tomato":90,
         "Potato":70,
         "Onion":45,
         "Carrot":120,
         "Cabbage":40,
         "Soyabean":80,
         "Pumpkin":40,
         "Squash":60,
         "Okara":90,
         "Onion":110,
         "Mushroom":180,
         "Sugarbeet":90,
         "Lettuce":60,
         "Celery":180,
         "Parsley":180,
         "Mint":180,
         "Turnip":60,
         "Tamarind":150,
         "Tofu":100,
         "Apple":120,
         "Lime":250,
         "Pomegranate":200,
         "Ginger":130,
         "Chilli":250,
         "Garlic":290,
         "Fish":270
        }

def totalcost(rams_list):
# YOUR CODE HERE
    return sum([price[i] for i in rams_list if i in price])

```

Output:
```
380
```

<!-- grade_id: cell-e631bda315d65fc4, points: 2, grade: true, solution: false -->
```python
rams_list =["Tomato","Garlic"]
assert(totalcost(rams_list)==380)
```

## Exercise 8: Dictionary

<b><div style="text-align: right">[POINTS: 3]</div></b>

**Task:** Write a function called `squares` that when given an integral number `n`, returns a dictionary that maps $i$ to $i^2$ such that $i$ is an integral number between 1 and `n` (both included).
For example, if `n` is 8, the `squares` should return this dictionary:
`{1: 1, 2: 4, 3: 9, 4: 16, 5: 25, 6: 36, 7: 49, 8: 64}`

<!-- grade_id: cell-6858c70d1575eb8c, points: 0, grade: false, solution: true -->
```python
def squares(n):
# YOUR CODE HERE
        return {n:n**2 for n in range(1,n+1)}
```

<!-- grade_id: cell-65df0d30b845126a, points: 3, grade: true, solution: false -->
```python
assert squares(1) is not None
```

## Exercise 9: Sets

<b><div style="text-align: right">[POINTS: 2]</div></b>

**Task:** Find the set difference of the sets `A` and `B`, store it in `ans` and remove the character `'e'` from the resulting set.


<!-- grade_id: cell-0c3d90a9f6014b64, points: 0, grade: false, solution: true -->
```python
B = {4, 5, 6, 7, 8, 'e'}
A = {1, 2, 3, 4, 5, 6}
ans = None
# YOUR CODE HERE
c = B.difference(A)
d = A.difference(B)
ans=c
ans.update(d)
ans.remove('e')
print(ans)


```

Output:
```
{1, 2, 3, 7, 8}
```

<!-- grade_id: cell-c1447e34ed140a05, points: 2, grade: true, solution: false -->
```python
assert ans is not None
```

## Exercise 10: Control Flow Statements

<b><div style="text-align: right">[POINTS: 2]</div></b>

**Task:** Write a program which will find all such numbers which are divisible by 7 but are not a multiple of 5,
between `2000` and `3200` (both included).
The numbers obtained should be stored in a list called `multiples_of_seven` in ascending order.

<!-- grade_id: cell-774d33765dd9e679, points: 0, grade: false, solution: true -->
```python
multiples_of_seven = []
# YOUR CODE HERE
for i in range(2000,3201):
    if i % 7 ==0 and i % 5 != 0:
        multiples_of_seven.append(i)
```

<!-- grade_id: cell-34e689983bc0abad, points: 2, grade: true, solution: false -->
```python
assert multiples_of_seven is not None
```

## Exercise 11: Hailstone Series

<b><div style="text-align: right">[POINTS: 3]</div></b>

The hailstone series is defined as follows: The $i^{th}$ value of this series, say $\text{H}[i]$, is equal to $\frac{\text{H}[i-1]}{2}$ if $\text{H}[i-1]$ is odd, and $3*\text{H}[i-1]+1$ if it is even. The hailstone series converges to 1 for any input integer greater than 1. While this is yet to be formally proved, experiments show that this holds for integers as large as $1.2 * 10^{12}$.

**Task:** Build a function `hailstone(n)` that takes in a number `n` and returns a list of the hailstone series starting at $n$. Your function should return the list when the last element is 1 for the first time. For example, if you start at 10, the returned list should be [10,5,16,8,4,2,1].

<!-- grade_id: cell-ddbcf0f7c936bbf4, points: 0, grade: false, solution: true -->
```python
def hailstone(n):
    numbers = [n]
    while n!=1:
        if n%2 == 0:
            n=n//2
            numbers.append(n)
        else:
            n=3*n+1
            numbers.append(n)
    return numbers
# YOUR CODE HERE
print(hailstone(10))
```

Output:
```
[10, 5, 16, 8, 4, 2, 1]
```

<!-- grade_id: cell-a9d789269089db25, points: 3, grade: true, solution: false -->
```python
assert (hailstone(10)==[10,5,16,8,4,2,1])
```

## Exercise 12: FInding the GCD

<b><div style="text-align: right">[POINTS: 3]</div></b>

**Task:** Using concept of looping, control flow statements and fucntions,create a function `gcd()` that takes in two positive integers `a` and `b` and returns their Greatest Common Divisor (GCD). 

<!-- grade_id: cell-a27f987e5613c584, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def gcd(a,b):
    while(b):
        a,b = b, a%b
    return a
```

<!-- grade_id: cell-93aba65a46bd9af6, points: 3, grade: true, solution: false -->
```python
# This shows the runtime of your function
# time.time() gives the number of seconds since the epoch as a float
import time
start_time = time.time()
a = 39203
b = 39400
print(f'GCD of {a} and {b} is {gcd(a,b)}.')
print(f'Runtime of gcd for a = {a} and b = {b} is {time.time() - start_time:.3f} seconds.')

```

Output:
```
GCD of 39203 and 39400 is 197.
Runtime of gcd for a = 39203 and b = 39400 is 0.000 seconds.
```

## Exercise 13: Reverse Strings

<b><div style="text-align: right">[POINTS: 1]</div></b>

**Task:** Find the reverse of string in `string` using concept of slicing and store it in a variable named `reverse_string`.

<!-- grade_id: cell-d82e4b5941680c86, points: 0, grade: false, solution: true -->
```python
string="I love programming in python."
reverse_string=None
# YOUR CODE HERE
reverse_string=string[::-1]
```

<!-- grade_id: cell-ebfdc5676218e490, points: 1, grade: true, solution: false -->
```python
assert reverse_string is not None
```

## Exercise 14: Matrix Multiplication

<b><div style="text-align: right">[POINTS: 3]</div></b>

**Task:**  Carry out matrix multiplication using list comprehension with a function `matmultiply`, that takes in two matrices (lists of lists) `A` and `B` and returns their matrix multiplication. The built-in `zip` function might be handy for this task.

<!-- grade_id: cell-26d2d26289fc2e1b, points: 0, grade: false, solution: true -->
```python
def matmultiply(A, B):
# YOUR CODE HERE
    zip_b = list(zip(*B))
    print(*B)
#     return [[sum(a*b for a,b in zip(row_a,col_b)) for col_b in zip_b] for row_a in A]
```

<!-- grade_id: cell-e93564ba74baaa74, points: 3, grade: true, solution: false -->
```python
A = [[1,2],[3,4]]
B = [[5,6],[7,8]]
assert(matmultiply(A,B)[0][0]==19)
```

## Exercise 15: Lambda Functions

<b><div style="text-align: right">[POINTS: 1]</div></b>

**Task:** Create a lambda function `exponent` to evaluate two times the exponential of a number `num` with base set to e. That is, exponent should evaluate the $2*e^x$ for a given number $x$. You can use `math.exp` for the exponential.

<!-- grade_id: cell-248a4e7a8edb3525, points: 0, grade: false, solution: true -->
```python
import math
exponent = lambda x : 2 * math.exp(x)
# YOUR CODE HERE
print(exponent(2))
```

Output:
```
14.7781121978613
```

<!-- grade_id: cell-a736e6da0b8a47b7, points: 1, grade: true, solution: false -->
```python
assert exponent is not None
```

## Exercise 16: Functional Programming

<b><div style="text-align: right">[POINTS: 2]</div></b>

**Task:** Using lambda functions, filter out the numbers divisible by 5 from the list `numbers` and store the resulting list in the variable `filtered_numbers`. Then, divide each number in the list `filtered_numbers` by 5 and then store this list in the variable `mapped_numbers`. 

<!-- grade_id: cell-030fc88fd141610b, points: 0, grade: false, solution: true -->
```python
numbers = [1, 5, 4, 6, 8, 11, 15, 3, 12, 75, 66, 99, 50, 100, 90, 555, 10257, 10725, 5, 10, 75, 76, 77, 78, 79, 80]
filtered_numbers = list(filter(lambda x:x%5==0,numbers))
doubled_numbers = list(map(lambda x:x//5,filtered_numbers))
# YOUR CODE HERE
print(filtered_numbers)
print(doubled_numbers)
```

Output:
```
[5, 15, 75, 50, 100, 90, 555, 10725, 5, 10, 75, 80]
[1, 3, 15, 10, 20, 18, 111, 2145, 1, 2, 15, 16]
```

<!-- grade_id: cell-963240d334c5615d, points: 2, grade: true, solution: false -->
```python
assert len(filtered_numbers) and len(doubled_numbers)
```

## Exercise 17: File Handling

<b><div style="text-align: right">[POINTS: 2]</div></b>

**Task:** Create a file named `'file.txt'` and write the strings in `line1`, `line2` and `line3` to the file **in order**. After writing, close the file.

<!-- grade_id: cell-570207d78f70e95b, points: 0, grade: false, solution: true -->
```python
line1='This is the first line\n'
line2='This is the second line\n'
line3='This is the third line\n'

# YOUR CODE HERE
with open('file.txt','w') as fp:
    fp.write(line1+ line2 +line3)
    
```

<!-- grade_id: cell-07abb3c6368c1860, points: 2, grade: true, solution: false -->
```python
### INTENTIONALLY LEFT BLANK
```

## Exercise 18: Reading From Files

<b><div style="text-align: right">[POINTS: 2]</div></b>

**Task:** Open the previously created file, read first 23 characters from the file into a variable `data` and store the current file position in variable named `position`. Then, close the file.

<!-- grade_id: cell-398e2a2da6daaf33, points: 0, grade: false, solution: true -->
```python
position = None
data = None
# YOUR CODE HERE
with open('file.txt','r') as fp:
    data = fp.read(23)
    position = fp.tell()
print(data)
print(position)
```

Output:
```
This is the first line

23
```

<!-- grade_id: cell-5cca88e8ec59fece, points: 2, grade: true, solution: false -->
```python
assert position is not None
assert data is not None

```

## Exercise 19: Parsing Words From File

<b><div style="text-align: right">[POINTS: 2]</div></b>

Reopen the file `file.txt` in read mode, read all the contents of the file and convert the string into list of words named `list_of_words`.

Hint: Use the `split` function on the string.

<!-- grade_id: cell-432ad9d3953c6b5c, points: 0, grade: false, solution: true -->
```python
list_of_words=None

# YOUR CODE HERE
with open('file.txt','r') as fp:
    list_of_words = fp.read().split('\n')
print(list_of_words)
```

Output:
```
['This is the first line', 'This is the second line', 'This is the third line', '']
```

<!-- grade_id: cell-29a28a5fc43dfd6b, points: 2, grade: true, solution: false -->
```python
assert list_of_words is not None

```

```python
# clean up the directory
!rm file.txt
```

```python

```