# Assignment 4.4 Similarity Measure Using Inversion

<b><div style="text-align: right">[POINTS: 40]</div></b>

Merge sort uses divide and conquer algorithm for the purpose of sorting. But it can be used to count numbers of inversions during the process as well. 

When sorting using merge sort (or any other sort algorithm), you can put the array you want to obtain on top and the array that you wish to sort at the bottom and then connect all the items respectivly. The number of intersections obtained will be you number of inversions. 

![inversions](https://docs.google.com/drawings/d/e/2PACX-1vTRGZzWkisPxRhiTQJuMiMV9i1yvPiniJt74JOwdk0fzdMD5tMFsIF-Kdtge5nIqnafojcIHH8K07f_/pub?w=585&h=281)

The number of inversions is the number of times elements of an array should be interchanged to be sorted. Hence, the number of inversions conducted during this process also signifies dissimilarity between these two arrays as how far (or close) the array is from being sorted. If an array is already sorted then number of inversions is 0 and if it is in reversed order, number of inversions is maximum i.e. `n*(n-1/2)`. For example, if our array is `[1,10,7,3,5]` then number of inversions will be 5. 

Now lets say you are given 10 items and you have to form an array using those items in such a way that the items have to be ordered in descending order according to your liking of the items. Your friend does the same. If you sort your friends array into your array and find number of inversions, you can evaluate dissimilarity between yours and your friend's choices.

Note: Even though this might not be the most efficient way to compute dissimilarity, many applications have used this technique.

### Exercise 1 : Merge Sort

<b><div style="text-align: right">[POINTS: 10]</div></b>

Given below are two arrays: items sorted by you and your friend. 

**Task:** Compute the dissimilarity of these choices as number of inversions by implementing merge sort using the function `merge_sort()` that accepts source array and dictionary containing values of each item to represent order of sorting. This function should return number of inversions.

<!-- grade_id: cell-15f1141fcb162d42, points: 0, grade: false, solution: true -->
```python
your_choice=["apple","cherry","banana","peanut","lime","pizza","icecream","guava","donut","cake"]
friend_choice=["apple","banana","guava","cherry","lime","icecream","peanut","donut","pizza","cake"]

target_dict={}
for index, item in enumerate(your_choice):
    target_dict[item]=index
    
def merge_sort(source,td,num_inversions=0): ## re
    '''
        The source is the list of choices of our friend and the td(target dict) is our dict of choices with indices
        
    '''
    if len(source) > 1:
        mid = len(source)//2
        left = source[:mid]
        #print(left)
        right = source[mid:]
        
        num_inversions = num_inversions + merge_sort(left, td)
        num_inversions = num_inversions + merge_sort(right,td)
        
        #Merging function 
        i = j = k = 0
        m, n = len(left), len(right)
        
        #Run the loop until the sum of indices of left and right becomes equal to the sum of their lengths
        while i + j < m + n:
            #When left is empty
            if i == m:
                source[k] = right[j]
                j, k = j+1, k+1
                
            #When right is empty 
            elif j == n: 
                source[k] = left[i]
                i, k = i+1, k+1
                
            #Left is smaller
            elif td[left[i]] <= td[right[j]]:
                source[k] = left[i]
                i, k = i+1, k+1
                
            #Right is smaller
            elif td[left[i]] > td[right[j]]:
                # This is because if an element in the left is greater than the one in right
                # Then all the elements next to it will also be greater than it
                # Since the arrays are sorted
                num_inversions += len(left)-i  
                source[k] = right[j]
                j, k = j+1, k+1
                
    return num_inversions
```

<!-- grade_id: cell-b9cab629d09a8da3, points: 10, grade: true, solution: false -->
```python
assert(merge_sort(["icecream","peanut","banana","lime","guava","donut","pizza","cake","apple","cherry"],target_dict)==23)
assert(merge_sort(["apple","cherry","banana","peanut","lime","pizza","icecream","guava","donut","cake"],target_dict)==0)

```

### Exercise 2 : Recommendation using Merge Sort
<b><div style="text-align: right">[POINTS: 10]</div></b>

Consider that you are an owner of a website. You have maintained an array of items for each person to store their likeness of items in descending order. For a given person A, you need to find number of inversions for all your users, find user B with least number of inversions according to A, and recommend items bought or liked by B to A.

Below is given choice of person A, and a dataset that consists of a list of list of items sorted by a person and items that he has bought. 

**Task:** Find the items that should be recommended to person A and store it in a list `recommended`. 

<!-- grade_id: cell-735cd44c20bb48c3, points: 0, grade: false, solution: true -->
```python
person_choice=["apple","banana","guava","cherry","lime","icecream","peanut","donut","pizza","cake"]
target_dict={}
for index, item in enumerate(person_choice):
    target_dict[item]=index
dataset=[[['icecream', 'cake', 'donut', 'apple', 'peanut', 'guava', 'cherry', 'lime', 'banana', 'pizza'], ['tuna', 'fish']], [['guava', 'donut', 'pizza', 'peanut', 'lime', 'icecream', 'cake', 'apple', 'banana', 'cherry'], ['shrimp', 'tuna']], [['donut', 'cake', 'cherry', 'lime', 'banana', 'pizza', 'peanut', 'icecream', 'guava', 'apple'], ['rice', 'burger']], [['lime', 'pizza', 'peanut', 'donut', 'cherry', 'guava', 'apple', 'banana', 'cake', 'icecream'], ['rice', 'cheese']], [['icecream', 'guava', 'lime', 'cherry', 'donut', 'pizza', 'banana', 'apple', 'cake', 'peanut'], ['fish', 'bread']], [['cake', 'cherry', 'apple', 'guava', 'lime', 'pizza', 'donut', 'banana', 'icecream', 'peanut'], ['sandwich', 'steak']], [['pizza', 'peanut', 'icecream', 'guava', 'lime', 'donut', 'banana', 'cake', 'apple', 'cherry'], ['fish', 'rice']], [['banana', 'donut', 'cake', 'lime', 'icecream', 'pizza', 'guava', 'apple', 'peanut', 'cherry'], ['steak', 'bread']], [['banana', 'icecream', 'donut', 'apple', 'guava', 'peanut', 'cake', 'pizza', 'cherry', 'lime'], ['sandwich', 'bread']], [['pizza', 'guava', 'peanut', 'banana', 'cherry', 'icecream', 'donut', 'lime', 'cake', 'apple'], ['tuna', 'bread']], [['lime', 'cake', 'peanut', 'pizza', 'donut', 'guava', 'apple', 'cherry', 'icecream', 'banana'], ['sandwich', 'shrimp']], [['apple', 'donut', 'icecream', 'guava', 'lime', 'banana', 'peanut', 'pizza', 'cherry', 'cake'], ['rice', 'fish']], [['cherry', 'guava', 'pizza', 'lime', 'donut', 'apple', 'peanut', 'icecream', 'cake', 'banana'], ['bread', 'burger']], [['donut', 'lime', 'pizza', 'cake', 'banana', 'guava', 'cherry', 'peanut', 'icecream', 'apple'], ['sandwich', 'tuna']], [['cake', 'peanut', 'guava', 'apple', 'banana', 'lime', 'icecream', 'donut', 'pizza', 'cherry'], ['shrimp', 'steak']], [['guava', 'cherry', 'cake', 'peanut', 'donut', 'lime', 'pizza', 'banana', 'apple', 'icecream'], ['burger', 'tuna']], [['cherry', 'pizza', 'peanut', 'banana', 'apple', 'cake', 'icecream', 'lime', 'donut', 'guava'], ['burger', 'rice']], [['guava', 'donut', 'peanut', 'cherry', 'banana', 'cake', 'lime', 'icecream', 'apple', 'pizza'], ['steak', 'bread']], [['peanut', 'apple', 'lime', 'cake', 'banana', 'pizza', 'donut', 'icecream', 'cherry', 'guava'], ['cheese', 'tuna']], [['icecream', 'donut', 'apple', 'peanut', 'pizza', 'cake', 'guava', 'banana', 'cherry', 'lime'], ['steak', 'bread']], [['lime', 'guava', 'cake', 'donut', 'apple', 'peanut', 'banana', 'pizza', 'icecream', 'cherry'], ['burger', 'rice']], [['pizza', 'guava', 'icecream', 'banana', 'apple', 'donut', 'cherry', 'cake', 'lime', 'peanut'], ['steak', 'fish']], [['apple', 'pizza', 'banana', 'cherry', 'donut', 'peanut', 'icecream', 'cake', 'lime', 'guava'], ['rice', 'cheese']], [['icecream', 'guava', 'apple', 'donut', 'peanut', 'cherry', 'lime', 'pizza', 'banana', 'cake'], ['burger', 'sandwich']], [['lime', 'donut', 'icecream', 'peanut', 'banana', 'pizza', 'cake', 'guava', 'apple', 'cherry'], ['cheese', 'rice']], [['guava', 'banana', 'pizza', 'lime', 'peanut', 'apple', 'icecream', 'cherry', 'cake', 'donut'], ['sandwich', 'shrimp']], [['cake', 'apple', 'icecream', 'lime', 'banana', 'cherry', 'peanut', 'donut', 'guava', 'pizza'], ['sandwich', 'bread']], [['guava', 'pizza', 'donut', 'apple', 'icecream', 'lime', 'cake', 'peanut', 'banana', 'cherry'], ['fish', 'steak']], [['apple', 'banana', 'guava', 'donut', 'lime', 'peanut', 'icecream', 'cherry', 'pizza', 'cake'], ['fish', 'steak']], [['banana', 'guava', 'cherry', 'peanut', 'icecream', 'cake', 'lime', 'apple', 'donut', 'pizza'], ['cheese', 'bread']], [['icecream', 'guava', 'banana', 'cherry', 'pizza', 'lime', 'donut', 'cake', 'peanut', 'apple'], ['rice', 'sandwich']], [['banana', 'donut', 'cherry', 'guava', 'lime', 'peanut', 'apple', 'icecream', 'cake', 'pizza'], ['fish', 'bread']], [['donut', 'lime', 'cherry', 'cake', 'guava', 'pizza', 'apple', 'peanut', 'icecream', 'banana'], ['shrimp', 'rice']], [['icecream', 'peanut', 'lime', 'banana', 'cherry', 'apple', 'donut', 'cake', 'guava', 'pizza'], ['rice', 'steak']], [['banana', 'donut', 'apple', 'cake', 'lime', 'icecream', 'peanut', 'pizza', 'guava', 'cherry'], ['rice', 'shrimp']], [['pizza', 'lime', 'cake', 'guava', 'banana', 'peanut', 'icecream', 'cherry', 'apple', 'donut'], ['steak', 'bread']], [['peanut', 'cherry', 'banana', 'apple', 'icecream', 'donut', 'lime', 'pizza', 'cake', 'guava'], ['sandwich', 'steak']], [['pizza', 'lime', 'guava', 'cake', 'apple', 'peanut', 'donut', 'icecream', 'banana', 'cherry'], ['shrimp', 'tuna']], [['guava', 'cherry', 'apple', 'icecream', 'peanut', 'donut', 'pizza', 'cake', 'banana', 'lime'], ['sandwich', 'burger']], [['guava', 'icecream', 'lime', 'banana', 'apple', 'pizza', 'peanut', 'donut', 'cake', 'cherry'], ['steak', 'bread']], [['pizza', 'apple', 'icecream', 'peanut', 'banana', 'donut', 'cake', 'guava', 'cherry', 'lime'], ['bread', 'shrimp']], [['cherry', 'lime', 'apple', 'pizza', 'donut', 'guava', 'banana', 'cake', 'peanut', 'icecream'], ['shrimp', 'rice']], [['banana', 'pizza', 'apple', 'lime', 'peanut', 'guava', 'icecream', 'cake', 'donut', 'cherry'], ['shrimp', 'rice']], [['lime', 'cake', 'pizza', 'banana', 'apple', 'icecream', 'donut', 'peanut', 'guava', 'cherry'], ['cheese', 'steak']], [['lime', 'peanut', 'icecream', 'cherry', 'pizza', 'banana', 'donut', 'apple', 'guava', 'cake'], ['rice', 'sandwich']], [['guava', 'pizza', 'lime', 'peanut', 'donut', 'cake', 'cherry', 'apple', 'banana', 'icecream'], ['sandwich', 'fish']], [['cake', 'guava', 'peanut', 'donut', 'apple', 'cherry', 'banana', 'icecream', 'lime', 'pizza'], ['steak', 'fish']], [['donut', 'peanut', 'banana', 'pizza', 'icecream', 'lime', 'cherry', 'cake', 'guava', 'apple'], ['shrimp', 'burger']], [['lime', 'donut', 'pizza', 'banana', 'cake', 'apple', 'cherry', 'guava', 'peanut', 'icecream'], ['rice', 'burger']], [['cherry', 'peanut', 'banana', 'lime', 'guava', 'apple', 'cake', 'pizza', 'icecream', 'donut'], ['sandwich', 'shrimp']]]

recommended =None
inversion_list = []
for r in dataset:
    inversion_list.append(merge_sort(r[0],target_dict))
    
least_inversion = min(inversion_list)
recommended = dataset[inversion_list.index(least_inversion)][1]

```

<!-- grade_id: cell-2662f8d71a351668, points: 10, grade: true, solution: false -->
```python
assert recommended is not None
```

### Exercise 3 : Sorting Dataset According to Number of Inversions
<b><div style="text-align: right">[POINTS: 10]</div></b>

**Task:** Add the value of number of inversions at first position in each row of the dataset, sort the dataset items in ascending order of number of inversions and store the result in variable `sorted_dataset`. Use any of the sort algorithms.

Note: If two items have same number of inversions, order them in lexicographical order.

<!-- grade_id: cell-f56ee68cc04337a5, points: 0, grade: false, solution: true -->
```python
person_choice=["apple","banana","guava","cherry","lime","icecream","peanut","donut","pizza","cake"]
dataset=[[['icecream', 'cake', 'donut', 'apple', 'peanut', 'guava', 'cherry', 'lime', 'banana', 'pizza'], ['tuna', 'fish']], [['guava', 'donut', 'pizza', 'peanut', 'lime', 'icecream', 'cake', 'apple', 'banana', 'cherry'], ['shrimp', 'tuna']], [['donut', 'cake', 'cherry', 'lime', 'banana', 'pizza', 'peanut', 'icecream', 'guava', 'apple'], ['rice', 'burger']], [['lime', 'pizza', 'peanut', 'donut', 'cherry', 'guava', 'apple', 'banana', 'cake', 'icecream'], ['rice', 'cheese']], [['icecream', 'guava', 'lime', 'cherry', 'donut', 'pizza', 'banana', 'apple', 'cake', 'peanut'], ['fish', 'bread']], [['cake', 'cherry', 'apple', 'guava', 'lime', 'pizza', 'donut', 'banana', 'icecream', 'peanut'], ['sandwich', 'steak']], [['pizza', 'peanut', 'icecream', 'guava', 'lime', 'donut', 'banana', 'cake', 'apple', 'cherry'], ['fish', 'rice']], [['banana', 'donut', 'cake', 'lime', 'icecream', 'pizza', 'guava', 'apple', 'peanut', 'cherry'], ['steak', 'bread']], [['banana', 'icecream', 'donut', 'apple', 'guava', 'peanut', 'cake', 'pizza', 'cherry', 'lime'], ['sandwich', 'bread']], [['pizza', 'guava', 'peanut', 'banana', 'cherry', 'icecream', 'donut', 'lime', 'cake', 'apple'], ['tuna', 'bread']], [['lime', 'cake', 'peanut', 'pizza', 'donut', 'guava', 'apple', 'cherry', 'icecream', 'banana'], ['sandwich', 'shrimp']], [['apple', 'donut', 'icecream', 'guava', 'lime', 'banana', 'peanut', 'pizza', 'cherry', 'cake'], ['rice', 'fish']], [['cherry', 'guava', 'pizza', 'lime', 'donut', 'apple', 'peanut', 'icecream', 'cake', 'banana'], ['bread', 'burger']], [['donut', 'lime', 'pizza', 'cake', 'banana', 'guava', 'cherry', 'peanut', 'icecream', 'apple'], ['sandwich', 'tuna']], [['cake', 'peanut', 'guava', 'apple', 'banana', 'lime', 'icecream', 'donut', 'pizza', 'cherry'], ['shrimp', 'steak']], [['guava', 'cherry', 'cake', 'peanut', 'donut', 'lime', 'pizza', 'banana', 'apple', 'icecream'], ['burger', 'tuna']], [['cherry', 'pizza', 'peanut', 'banana', 'apple', 'cake', 'icecream', 'lime', 'donut', 'guava'], ['burger', 'rice']], [['guava', 'donut', 'peanut', 'cherry', 'banana', 'cake', 'lime', 'icecream', 'apple', 'pizza'], ['steak', 'bread']], [['peanut', 'apple', 'lime', 'cake', 'banana', 'pizza', 'donut', 'icecream', 'cherry', 'guava'], ['cheese', 'tuna']], [['icecream', 'donut', 'apple', 'peanut', 'pizza', 'cake', 'guava', 'banana', 'cherry', 'lime'], ['steak', 'bread']], [['lime', 'guava', 'cake', 'donut', 'apple', 'peanut', 'banana', 'pizza', 'icecream', 'cherry'], ['burger', 'rice']], [['pizza', 'guava', 'icecream', 'banana', 'apple', 'donut', 'cherry', 'cake', 'lime', 'peanut'], ['steak', 'fish']], [['apple', 'pizza', 'banana', 'cherry', 'donut', 'peanut', 'icecream', 'cake', 'lime', 'guava'], ['rice', 'cheese']], [['icecream', 'guava', 'apple', 'donut', 'peanut', 'cherry', 'lime', 'pizza', 'banana', 'cake'], ['burger', 'sandwich']], [['lime', 'donut', 'icecream', 'peanut', 'banana', 'pizza', 'cake', 'guava', 'apple', 'cherry'], ['cheese', 'rice']], [['guava', 'banana', 'pizza', 'lime', 'peanut', 'apple', 'icecream', 'cherry', 'cake', 'donut'], ['sandwich', 'shrimp']], [['cake', 'apple', 'icecream', 'lime', 'banana', 'cherry', 'peanut', 'donut', 'guava', 'pizza'], ['sandwich', 'bread']], [['guava', 'pizza', 'donut', 'apple', 'icecream', 'lime', 'cake', 'peanut', 'banana', 'cherry'], ['fish', 'steak']], [['apple', 'banana', 'guava', 'donut', 'lime', 'peanut', 'icecream', 'cherry', 'pizza', 'cake'], ['fish', 'steak']], [['banana', 'guava', 'cherry', 'peanut', 'icecream', 'cake', 'lime', 'apple', 'donut', 'pizza'], ['cheese', 'bread']], [['icecream', 'guava', 'banana', 'cherry', 'pizza', 'lime', 'donut', 'cake', 'peanut', 'apple'], ['rice', 'sandwich']], [['banana', 'donut', 'cherry', 'guava', 'lime', 'peanut', 'apple', 'icecream', 'cake', 'pizza'], ['fish', 'bread']], [['donut', 'lime', 'cherry', 'cake', 'guava', 'pizza', 'apple', 'peanut', 'icecream', 'banana'], ['shrimp', 'rice']], [['icecream', 'peanut', 'lime', 'banana', 'cherry', 'apple', 'donut', 'cake', 'guava', 'pizza'], ['rice', 'steak']], [['banana', 'donut', 'apple', 'cake', 'lime', 'icecream', 'peanut', 'pizza', 'guava', 'cherry'], ['rice', 'shrimp']], [['pizza', 'lime', 'cake', 'guava', 'banana', 'peanut', 'icecream', 'cherry', 'apple', 'donut'], ['steak', 'bread']], [['peanut', 'cherry', 'banana', 'apple', 'icecream', 'donut', 'lime', 'pizza', 'cake', 'guava'], ['sandwich', 'steak']], [['pizza', 'lime', 'guava', 'cake', 'apple', 'peanut', 'donut', 'icecream', 'banana', 'cherry'], ['shrimp', 'tuna']], [['guava', 'cherry', 'apple', 'icecream', 'peanut', 'donut', 'pizza', 'cake', 'banana', 'lime'], ['sandwich', 'burger']], [['guava', 'icecream', 'lime', 'banana', 'apple', 'pizza', 'peanut', 'donut', 'cake', 'cherry'], ['steak', 'bread']], [['pizza', 'apple', 'icecream', 'peanut', 'banana', 'donut', 'cake', 'guava', 'cherry', 'lime'], ['bread', 'shrimp']], [['cherry', 'lime', 'apple', 'pizza', 'donut', 'guava', 'banana', 'cake', 'peanut', 'icecream'], ['shrimp', 'rice']], [['banana', 'pizza', 'apple', 'lime', 'peanut', 'guava', 'icecream', 'cake', 'donut', 'cherry'], ['shrimp', 'rice']], [['lime', 'cake', 'pizza', 'banana', 'apple', 'icecream', 'donut', 'peanut', 'guava', 'cherry'], ['cheese', 'steak']], [['lime', 'peanut', 'icecream', 'cherry', 'pizza', 'banana', 'donut', 'apple', 'guava', 'cake'], ['rice', 'sandwich']], [['guava', 'pizza', 'lime', 'peanut', 'donut', 'cake', 'cherry', 'apple', 'banana', 'icecream'], ['sandwich', 'fish']], [['cake', 'guava', 'peanut', 'donut', 'apple', 'cherry', 'banana', 'icecream', 'lime', 'pizza'], ['steak', 'fish']], [['donut', 'peanut', 'banana', 'pizza', 'icecream', 'lime', 'cherry', 'cake', 'guava', 'apple'], ['shrimp', 'burger']], [['lime', 'donut', 'pizza', 'banana', 'cake', 'apple', 'cherry', 'guava', 'peanut', 'icecream'], ['rice', 'burger']], [['cherry', 'peanut', 'banana', 'lime', 'guava', 'apple', 'cake', 'pizza', 'icecream', 'donut'], ['sandwich', 'shrimp']]]

new_dataset = dataset
i = 0
for row in new_dataset:
    row.insert(0, inversion_list[i])
    i += 1
    
new_dataset.sort() 
sorted_dataset = new_dataset
```

<!-- grade_id: cell-93aab7b66f928688, points: 10, grade: true, solution: false -->
```python
assert sorted_dataset is not None
```

### Exercise 4 : Binary Search
<b><div style="text-align: right">[POINTS: 10]</div></b>

Often you will need to find an item for given number of inversions. Search will be faster if you implement binary search on the sorted dataset than using linear search. 

**Task:** Write a function `binary_search()` that takes the sorted dataset, lowest index to search for, highest index to search for and value of number of inversions to search index at which item having `n` inversions for person A lies. Also return the corresponding items bought by that person and total steps involved during the search.

Note: use `mid = left + (right - l)//2` to find the location of mid value. Also since there exists same number of inversions for different items, tests will be conducted only for unique values.

<!-- grade_id: cell-17f6e912e8c36a55, points: 0, grade: false, solution: true -->
```python
def binary_search(arr, left, right, n): 
    count = 0
    while left <= right:
        count = count + 1
        mid = left + (right - left)//2;
        if arr[mid][0] == n:
            return mid, arr[mid][2], count
        elif arr[mid][0] < n:
            left = mid + 1
        else:
            right = mid -1   
    return -1
binary_search(sorted_dataset, 0,len(sorted_dataset)-1, 8)
```

Output:
```
(0, ['fish', 'steak'], 5)
```

<!-- grade_id: cell-f013f58677f5d9b7, points: 10, grade: true, solution: false -->
```python
assert (binary_search(sorted_dataset, 0,len(sorted_dataset)-1, 8)==(0, ['fish', 'steak'], 5))
```

```python

```