# Algorithmic Challenges

<b><div style="text-align: right">[TOTAL POINTS: 50]</div></b>

In this assignment, we will explore some of challenges or problems that can be solved using basic algorithms and data scructures such as stack, queue and linked list.

## Problem 1: Balanced Parenthesis

Stacks can be used to check whether the given expression has balanced symbols. This algorithm is very useful in compilers. Each time the parser parses one character at a time. If the character is an opening delimiter such as `(, [` or `{` then it is written to the stack. When a closing delimiter is encountered like `),]` or `}`, the stack is popped. 
The opening and closing delimeters are then compared. If they match, the parsing of the string continues. If they do not match, the parser indicates that there is an error on the line. A linear-time 0(n) algorithm based on stack can be given as:

**Algorithm:** 

1. Create a stack. 
2. while (end of input is not reached):
    1. If the character is an opening symbol like `(, [ , {` push it onto the stack. 
    2. If it is a closing symbol like `), ], }`, then if the stack is empty report an error. Otherwise pop the stack. 
    3. If the symbol popped is not the corresponding opening symbol, report an error. 
3. At end of input. if the stack is not empty report an error 

In python, you can use list as well as string operations. For list, you can use the `.append()` and `.pop()` method.

### Exercise 1: Check Balanced Parenthesis

<b><div style="text-align: right">[POINTS: 10]</div></b>

**Task:** Implement an algorithm to check if an expression has balanced parenthesis.

**Note** Name of the attributes **should** be `length` and `breadth`.

<!-- grade_id: cell-21356d6c2a5c1527, points: 0, grade: false, solution: true -->
```python
parent = {
    '}' :'{',
    ']':'[',
    ')' :'('
}

def check_balanced_parenthesis(expression):
    # YOUR CODE HERE
    stack=[]
    for symbol in expression:
        if symbol in '({[':
            stack.append(symbol)
        else:
            if len(stack) == 0:
                return False
            top_symbol = stack.pop()
            if parent[symbol] != top_symbol:
                return False
    if len(stack) != 0:
            return False
    return True

```

<!-- grade_id: cell-c967e01b8e918a4d, points: 10, grade: true, solution: false -->
```python
assert check_balanced_parenthesis('{[()]}')
assert check_balanced_parenthesis('([{}])')

```

## Queue: Down to Zero


**Down To Zero:**

Let us suppose you have $K$ element Array. Each element is a single integer number $N$. For each element you can perform any of the two operations:

1: If $a$ and $b$ are two integers such that: $N = a \times b$, where $a\neq 1$ and $b\neq 1$, replace $N$ by $\text{max}(a,b)$

2: Decrease the value of $N$ by 1, i.e. $N=N-1$

You have to determine the minimum number of moves required to reduce the value of $N$ to $0$.

### Exercise 2: Down to Zero Implementation

<b><div style="text-align: right">[POINTS: 20]</div></b>

**Task:**

- Create a class `DownToZero` that stores the list of minimum moves for each number corresponding to its index.

- Create a method `best_factors` that can find the next value of $N$ according to operation (1). You can find the factors that will lead to least number of moves.

- Create a function `minimum_steps` that can find the minimum number of moves using either operation (1) or (2).

- Create a function `get_number_of_steps` that returns minimum number of moves required by given values of arrays

**How to structure your code ?**

- First, specify that the minimum number of moves is 0 for $N=0$ in a answer list.
- Second, ranging from 1 to `max_range` find the minimum number of steps required to get to zero and add it to answer list.

***Tips:*** Use previously found minimum number to find minimum number for larger numbers.

<!-- grade_id: cell-4f7ae1b2f20d1923, points: 0, grade: false, solution: true -->
```python
import math
class DownToZero(object):
    
    def __init__(self,max_range = 1000):
        self.max_range = max_range
        self.data_list = [0,]
        for i in range(1,self.max_range):
            self.down_one_step(i)
    
    def best_factors(self,x):
        factors = set()
        for i in range(2, (x//2)+1):
            if x%i == 0:
                y = x//i
                factors.add(max(i,y))
        return min(factors)
    
    def down_one_step(self, value):
        if value == 0: return 0
        Q = [ (value,0) ]
        setQ = [ 0 ] * value
        while Q:
            value, steps = Q.pop(0)
            if value == 1: return steps+1
            div = int(math.sqrt( value ))
            while div > 1:
                if value % div == 0 and not setQ[value // div]:
                    Q.append( (value // div, steps+1) )
                    setQ[ value // div ] = 1
                div -= 1
                
            if not setQ[value-1]:
                Q.append( (value-1, steps+1) )
                setQ[ value-1 ] = 1
    
    def minimum_steps(self, value):
        return self.down_on_step(value)
        
    
    def get_number_of_steps(self, values:list):
        answer = []
        for i in values:
            answer.append(self.down_one_step(i))
        return answer

```

<!-- grade_id: cell-ef38a5dbb5869d1a, points: 20, grade: true, solution: false -->
```python
### INTENTIONALLY LEFT BLANK

```

### Exercise 3: Doubly Linked List : Delete at an Arbitrary Index


<b><div style="text-align: right">[POINTS: 20]</div></b>

**Task:**

Doubly Linked List is a type of Linked List where we can traverse the nodes in each direction. We have previously written a `DoublyLinkedList` class with methods to `traverse`,`insert_at_beg`,`insert_at_end`, `insert_at_index`, `delete_at_beg`, `delete_at_end`, `delete_at_index` and `search`.

Write a method for `DoublyLinkedList` class to delete element at given index, i.e. `delete_at_index`.

```python
class Node:
    #initializer
    def __init__(self,data = None):
        self.data = data
        self.next = None
        self.prev = None
    
    def __str__(self):
        return str(data)   
```

<!-- grade_id: cell-483d4c597c29dcbb, points: 0, grade: false, solution: true -->
```python
class DoublyLinkedList:
    def __init__(self):
        self.head = None

    ############################Traverse#############################
    def traverse(self):
        if self.head is None:
            print("The list is empty")
            return
        else:
            n = self.head
            while n is not None:
                print(n.data , " ")
                n = n.next

    ############################INSERT#############################
    def insert_at_beg(self, data):
        new_node = Node(data)
        if self.head is None:
            self.head = new_node
            return
        new_node.next = self.head
        self.head.prev = new_node
        self.head = new_node


    def insert_at_end(self, data):
        new_node = Node(data)
        if self.head is None:
            self.head = new_node
            return
        n = self.head
        while n.next is not None:
            n = n.next
        n.next = new_node
        new_node.prev = n


    def insert_at_index(self, index, data):
        if index == 1:
            #equivalent to inserting at beginning
            new_node = Node(data)
            if self.head is None:
                self.head = new_node
                return
            new_node.next = self.head
            self.head.prev = new_node
            self.head = new_node

        #traverse i=index nodes
        i = 1
        n = self.head
        while i < index-1 and n is not None:
            n = n.next
            i += 1
        if n is None:
            print("Index out of bound")
        else:
            new_node = Node(data)
            new_node.next = n.next
            new_node.prev = n
            n.next = new_node

    ############################INSERT#############################
    ############################DELETE#############################
    def delete_at_beg(self):
        if self.head is None:
            print("The list is empty")
            return
        if self.head.next is None:
            self.head = None
            return
        self.head = self.head.next
        # update previous pointer of new head to None
        self.head.prev = None


    def delete_at_end(self):
        if self.head is None:
            print("The list is empty")
            return
        if self.head.next is None:
            self.head = None
            return
        n = self.head
        while n.next is not None:
            n = n.next
        n.prev.next = None

    ########## Delete at index -- the assignment ##################

    def delete_at_index(self, index):
        if self.head is None:
            print("The list is empty")
            return
        j = self.head
        count = 1
        while j.next is not None:
            j = j.next
            count += 1
        if index < 0:
            print("Enter a valid index")
            return
        elif index == 0:
            self.delete_at_beg()
        elif index == count-1:
            self.delete_at_end()
        elif index < count:
            n = self.head
            i = 0
            while i < index-1 and n is not None:
                n = n.next
                i += 1
            if n is None:
                print("Index out of bound")
            else:
                n.next.next.prev = n
                n.next = n.next.next
        else:
            print('Index out of bound')

    ############################DELETE#############################
    ############################SEARCH#############################

    def search(self,data):
        if self.head is None:
            print("The list is empty")
            return
        else:
            n = self.head
            index = 1
            while n is not None:
                if n.data == data:
                    print("Data found at index ",index)
                    return True
                n = n.next
                index += 1
            print("Data not found")
            return False

d = DoublyLinkedList()
d.insert_at_beg(5)
d.insert_at_end(51)
d.insert_at_end(10)
d.delete_at_index(1)
d.traverse()

```

Output:
```
5  
10
```

<!-- grade_id: cell-163569573de5a621, points: 20, grade: true, solution: false -->
```python
### INTENTIONALLY LEFT BLANK

```

```python

```